```python
# Initialize Otter
import otter
grader = otter.Notebook("Test_assign.ipynb")
```

## DSA 2024 Summer School Admittance Check

Thanks for your interest in attending DSA 2024 Nyeri, Kenya. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Nyeri!

```python
import pandas as pd
import numpy as np

%matplotlib inline

grader = otter.Notebook()
```

Output:
```
File ~\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.11_qbz5n2kfra8p0\LocalCache\local-packages\Python311\site-packages\IPython\core\interactiveshell.py:3665, in InteractiveShell.enable_matplotlib(self, gui)
   3662     import matplotlib_inline.backend_inline
   3664 from IPython.core import pylabtools as pt
-> 3665 gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3667 if gui != None:
   3668     # If we have our first gui selection, store it
   3669     if self.pylab_gui_select is None:

File ~\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.11_qbz5n2kfra8p0\LocalCache\local-packages\Python311\site-packages\IPython\core\pylabtools.py:338, in find_gui_and_backend(gui, gui_select)
    321 def find_gui_and_backend(gui=None, gui_select=None):
    322     """Given a gui string return the gui and mpl backend.
    323 
    324     Parameters
   (...)
    335     'WXAgg','Qt4Agg','module://matplotlib_inline.backend_inline','agg').
    336     """
--> 3
```

**Question 1:** write a function `isValid(s)` that takes as argument a string s containing a sequence of parenthesis '(', ')', '{', '}', '[' and ']', and  determines if the input is valid. A input string is valid if for every open parenthensis there is a close one and parenthesis is well-formed. e.g  "(){}[]" is valid.

```python
def isValid(s):
    stack = []
    mapping = {')': '(', '}': '{', ']': '['}
    
    for char in s:
        if char in mapping.values():  # If it's an opening bracket
            stack.append(char)
        elif char in mapping.keys():  # If it's a closing bracket
            if stack == [] or mapping[char] != stack.pop():
                return False
        else:
            return False
    
    return stack == []

# Example usage:
print(isValid("(){}[]"))  # Output: True
print(isValid("([{}])"))  # Output: True
print(isValid("({[}])"))  # Output: False

```

Output:
```
True
True
False
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Given a paragraph as a string, write a function that return the number of character with odd frequencies. E.g The paragraph ``DSA 2024 Nyeri`` has *10* characters with odd frequencies. i.e the entire frequency count is given as {' ': 2, '2': 2, 'D': 1, 'S': 1, 'A': 1, '0': 1, '4': 1, 'N': 1, 'y': 1, 'e': 1, 'r': 1, 'i': 1}) and there are *10* characters with odd frequences. So the function should return *10*.

```python
def oddFrequencyCounter(theParagraph):
    # Step 1: Count the frequency of each character
    frequency = {}
    for char in theParagraph:
        if char in frequency:
            frequency[char] += 1
        else:
            frequency[char] = 1
    
    # Step 2: Count characters with odd frequencies
    odd_count = 0
    for count in frequency.values():
        if count % 2 != 0:
            odd_count += 1
    
    return odd_count

# Example usage:
paragraph = "DSA 2024 Nyeri"
print(oddFrequencyCounter(paragraph))  # Output: 10

    
```

Output:
```
10
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** Write an infinite generator function `odd_squares_sum` that yields the sum of square of odd numbers. e.g $1^2 + 3^2 + 5^2 + ...$ up to a ``limit``

```python
def odd_squares_sum(limit):
    current_sum = 0
    odd_number = 1
    
    while True:
        square = odd_number ** 2
        if square > limit:
            break
        current_sum += square
        yield current_sum
        odd_number += 2  # Move to the next odd number

# Example usage:
gen = odd_squares_sum(100)
for value in gen:
    print(value)

```

Output:
```
1
10
35
84
165
```

```python
grader.check("q3")
```

Output:
```
q3 results:
    q3 - 1 result:
        ❌ Test case failed
        Trying:
            points: 4
        Expecting nothing
        ok
        Trying:
            theGen = odd_squares_sum(12)
        Expecting nothing
        ok
        Trying:
            assert next(theGen) == 1
        Expecting nothing
        **********************************************************************
        Line 3, in q3 0
        Failed example:
            assert next(theGen) == 1
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.11_3.11.2544.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q3 0[2]>", line 1, in <module>
                assert next(theGen) == 1
                       ^^^^^^^^^^^^^^^^^
            AssertionError
        Trying:
            assert next(theGen) == 9
        Expectin
```

**Question 4:** Using the `odd_squares_sum` generator defined above, create a list of sum of squares up to a limit of $20$ and store the results in a numpy.array variable called `oddSumList`

```python
def odd_squares_sum(limit):
    current_sum = 0
    odd_number = 1
    
    while True:
        square = odd_number ** 2
        if square > limit:
            break
        current_sum += square
        yield current_sum
        odd_number += 2  # Move to the next odd number

# Example usage to create a list and then convert to numpy array
limit = 100
gen = odd_squares_sum(limit)

# Collect the results into a list
results = list(gen)

# Convert the list to a numpy array
oddSumList = np.array(results)

print(oddSumList)
```

Output:
```
[  1  10  35  84 165]
```

```python
grader.check("q4")
```

Output:
```
q4 results:
    q4 - 1 result:
        ❌ Test case failed
        Trying:
            points: 2
        Expecting nothing
        ok
        Trying:
            assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
        Expecting nothing
        **********************************************************************
        Line 2, in q4 0
        Failed example:
            assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.11_3.11.2544.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q4 0[1]>", line 1, in <module>
                assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
                        ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
```

**Question 5:** Compute the element-wise remainder of ``oddSumList`` when divided by $5$ and merge it with ``oddSumList``. The final output stored in the variable `mergedList` should be in the form of a list of tupples e.g ``[(1,1), (4,9), (0,25), ...]``

```python
def odd_squares_sum(limit):
    odd_numbers = [x for x in range(1, limit + 1) if x % 2 != 0]
    odd_squares = [(x, x ** 2) for x in odd_numbers]
    merged_list = [(x, x % limit) for x, y in odd_squares]
    return merged_list

# Example usage:
limit = 5
mergedList = odd_squares_sum(limit)
print(mergedList)

```

Output:
```
[(1, 1), (3, 3), (5, 0)]
```

```python
grader.check("q5")
```

Output:
```
q5 results:
    q5 - 1 result:
        ❌ Test case failed
        Trying:
            points: 2
        Expecting nothing
        ok
        Trying:
            assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
        Expecting nothing
        **********************************************************************
        Line 2, in q5 0
        Failed example:
            assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.11_3.11.2544.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q5 0[1]>", line 1, in <module>
                assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    while b:
        a, b = b, a % b
    return a

# Example usage:
print(greatest_common_divisor(10, 15))  # Output: 5
print(greatest_common_divisor(21, 14))  # Output: 7
print(greatest_common_divisor(100, 25)) # Output: 25

```

Output:
```
5
7
25
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:**  Write a function `get_3_nearest` that takes in a point of interest ``pt`` and a **list** of points ``ptlist``  and returns a list of 3 nearest points from the point of interest ``pt``. Assume the distance between any two point is defined by the `L1-norm`.

```python
def get_3_nearest(pt, ptlist):
    # Calculate the L1-norm (Manhattan distance) between two points
    def l1_norm(p1, p2):
        return sum(abs(a - b) for a, b in zip(p1, p2))
    
    # Calculate distances from pt to all points in ptlist
    distances = [(l1_norm(pt, other_pt), other_pt) for other_pt in ptlist]
    
    # Sort points by their distances
    distances.sort(key=lambda x: x[0])
    
    # Get the three nearest points
    nearest_points = [point for _, point in distances[:3]]
    
    return nearest_points

# Example usage:
pt = (1, 2)
ptlist = [(2, 3), (4, 5), (1, 1), (3, 3), (0, 0)]
print(get_3_nearest(pt, ptlist))  # Output: [(2, 3), (1, 1), (3, 3)]

```

Output:
```
[(1, 1), (2, 3), (3, 3)]
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

**Question 8:**  Write a function `diagonal_vector(M)` that returns a **numpy** array of the list of **absolute** values of the main diagonal entries in the matrix $M$

```python
def diagonal_vector(M):
    # Ensure M is a numpy array
    M = np.array(M)
    
    # Extract the main diagonal entries
    main_diagonal = np.diag(M)
    
    # Take absolute values
    abs_main_diagonal = np.abs(main_diagonal)
    
    return abs_main_diagonal

# Example usage:
M = np.array([[1, -2, 3],
              [4, 5, -6],
              [-7, 8, 9]])

result = diagonal_vector(M)
print(result)  # Output: [1 5 9]

```

Output:
```
[1 5 9]
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Question 9:**  Write a function `flatten_reverse_lists` that takes in a list of lists and outputs a **reverse** sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_reverse_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_reverse_lists(superlist):
    # Step 1: Flatten the list of lists into a single list
    flat_list = [item for sublist in superlist for item in sublist]
    
    # Step 2: Sort the flat list in reverse order
    sorted_list = sorted(flat_list, reverse=True)
    
    # Step 3: Return the sorted list
    return sorted_list

# Example usage:
superlist = [[2, 13, 44], [6, 7]]
result = flatten_reverse_lists(superlist)
print(result)  # Output: [44, 13, 7, 6, 2]
```

Output:
```
[44, 13, 7, 6, 2]
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

**Question 9:** Create a DataFrame mirroring the table below and assign this to `data`.

| flavor | scoops | price |
|-----|-----|-----|
| white chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| dark chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| white chocolate | 3 | 2 |
| dark chocolate | 3 | 3 |
| white chocolate | 2 | 2 |
| dark chocolate | 5 | 3 |


```python
data = [
    {'flavor': 'white chocolate', 'scoops': 1, 'price': 2},
    {'flavor': 'vanilla', 'scoops': 1, 'price': 1.5},
    {'flavor': 'dark chocolate', 'scoops': 2, 'price': 3},
    {'flavor': 'strawberry', 'scoops': 1, 'price': 2},
    {'flavor': 'strawberry', 'scoops': 3, 'price': 4},
    {'flavor': 'vanilla', 'scoops': 2, 'price': 2},
    {'flavor': 'mint', 'scoops': 1, 'price': 4},
    {'flavor': 'mint', 'scoops': 2, 'price': 5},
    {'flavor': 'white chocolate', 'scoops': 3, 'price': 2},
    {'flavor': 'dark chocolate', 'scoops': 3, 'price': 3},
    {'flavor': 'white chocolate', 'scoops': 2, 'price': 2},
    {'flavor': 'dark chocolate', 'scoops': 5, 'price': 3}
]

# Create the DataFrame
df = pd.DataFrame(data)

# Print the DataFrame
print(df)

```

Output:
```
flavor  scoops  price
0   white chocolate       1    2.0
1           vanilla       1    1.5
2    dark chocolate       2    3.0
3        strawberry       1    2.0
4        strawberry       3    4.0
5           vanilla       2    2.0
6              mint       1    4.0
7              mint       2    5.0
8   white chocolate       3    2.0
9    dark chocolate       3    3.0
10  white chocolate       2    2.0
11   dark chocolate       5    3.0
```

**Question 10:** Do the following to the dataframe:
* Create a new collumn ``total_price`` whose value is equal to $scoops * price$*
* Write a function ``groupStatistics(data, groupValue)``. Internally, this function groups ``data``  by ``flavor`` and then returns statistics of a given grouped item ``groupValue`` indexed on the ``total_price`` columns. The statistics is a numpy array contains ``[mean, media, min, max, std]`` of the ``total_price`` column. The ``std`` should be rounded to 2 **decimal places**



```python
df['total_price'] = df['scoops'] * df['price']

# Step 2: Define the function groupStatistics
def groupStatistics(data, groupValue):
    # Group by flavor
    grouped = data.groupby('flavor')
    
    # Filter based on groupValue
    group_data = grouped.get_group(groupValue)
    
    # Calculate statistics on 'total_price' column
    mean = group_data['total_price'].mean()
    median = group_data['total_price'].median()
    minimum = group_data['total_price'].min()
    maximum = group_data['total_price'].max()
    std = round(group_data['total_price'].std(), 2)
    
    # Return statistics as a numpy array
    return np.array([mean, median, minimum, maximum, std])

# Example usage:
groupValue = 'white chocolate'
statistics = groupStatistics(df, groupValue)
print(f"Statistics for {groupValue}: {statistics}")
```

Output:
```
Statistics for white chocolate: [4. 4. 2. 6. 2.]
```

```python
grader.check("q11")
```

Output:
```
q11 results:
    q11 - 1 result:
        ❌ Test case failed
        Trying:
            points: 1
        Expecting nothing
        ok
        Trying:
            np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
        Expecting nothing
        **********************************************************************
        Line 2, in q11 0
        Failed example:
            np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.11_3.11.2544.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q11 0[1]>", line 1, in <module>
                np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
                                
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the exported ZIP. Take note of the ZIP number and proceed to fill the summer school form

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
Failed to start the Kernel. 

Unable to start Kernel 'Python 3.12.2' due to a timeout waiting for the ports to get used. 

View Jupyter <a href='command:jupyter.viewOutput'>log</a> for further details.
```

Tests for q1 (points: 4)
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('{[()]}}') == False
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('') == True
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('(){}[]') == True
>>> test_isValid(isValid)
```


Tests for q11 (points: 5)
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'dark chocolate'), [10.0, 9.0, 6.0, 4.58])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'white chocolate'), [4.0, 4.0, 2.0, 2.0])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'vanilla'), [2.75, 2.75, 1.5, 1.77])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'mint'), [7.0, 7.0, 4.0, 4.24])
```


Tests for q2 (points: 4)
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'DSA 2024 Nyeri'
...     assert oddFrequencyCounter(paragraph) == 10
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = "The following is sample text I used to practice special characters using keybr.com: 112233445566778899 Saturn V rocket's first stage carries 203,400 gallons (770,000 liters) of kerosene fuel and 318,000 gallons (1.2 million liters) of liquid oxygen needed for combustion. At liftoff, the stage's five F-1 rocket engines ignite and produce 7.5 million pounds of thrust. To replace those goofy quantities with the far less retarded metric system (even though liters are considered part of the metric system they are the same as cubic deci-meters) you would say 770 cubic meters of kerosene {abbreviated as m3} and 1,204 m3 of liquid O2 [O2 is the symbol for oxygen]. We would also say it produced 33,600,000 newtons of force <abbreviated as N>. To add scientific notation {a way of writing numbers that allows you to write only as many digits of specificity as you would like} you can write 7.7 * 10 ^ 2 m3 of kerosene 1.204 * 10 ^ 3 m3 of O2 and 3.3 * 10 ^ 7 newtons."
...     assert oddFrequencyCounter(paragraph) == 40
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'Data Science Africa in collaboration with The Swiss Federal Department of Foreign Affairs (FDFA), ETH Zurich, EPFL, the European Laboratory for Learning and Intelligent Systems (ELLIS), the Swiss National Computing Centre (CSCS), and the LUMI consortium launched the International Computation and AI Network (ICAIN) at the World Economic Forum (WEF) 2024 in Davos. Its mission is to develop AI technologies that benefit society as a whole, as well as being accessible to all and sustainable, thereby helping to reduce global inequality.'
...     assert oddFrequencyCounter(paragraph) == 20
```


Tests for q3 (points: 4)
```
>>> points: 4
>>> theGen = odd_squares_sum(12)
>>> assert next(theGen) == 1
>>> assert next(theGen) == 9
>>> assert next(theGen) == 25
>>> assert next(theGen) == 49
>>> assert next(theGen) == 81
>>> assert next(theGen) == 121
>>> del theGen
```


Tests for q4 (points: 4)
```
>>> points: 2
>>> assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
```
```
>>> points: 2
>>> assert isinstance(oddSumList, np.ndarray)
```


Tests for q5 (points: 2)
```
>>> points: 2
>>> assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
```


Tests for q6 (points: 5)
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(10, 15), 5)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(15, 19), 1)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(100, 105), 5)
```


Tests for q7 (points: 8)
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((12, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(13, 9), (10, 12), (5, 9)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((5, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(5, 9), (2, 4), (13, 9)])
```


Tests for q8 (points: 4)
```
>>> points: 2
>>> np.testing.assert_array_equal(diagonal_vector([[1, 2, 3, 4, 5, 6, 7, 8], [9, 10, 11, 12, 13, 14, 15, 16], [17, 18, 19, 20, 21, 22, 23, 24], [25, 26, 27, 28, 29, 30, 31, 32], [33, 34, 35, 36, 37, 38, 39, 40], [41, 42, 43, 44, 45, 46, 47, 48], [49, 50, 51, 52, 53, 54, 55, 56], [57, 58, 59, 60, 61, 62, 63, 64]]), [1, 10, 19, 28, 37, 46, 55, 64])
```
```
>>> points: 2
>>> assert isinstance(diagonal_vector([[1, 2], [3, 4]]), np.ndarray)
```


Tests for q9 (points: 3)
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2, 13, 44], [6, 7]]), [44, 13, 7, 6, 2])
```
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2], [61, 34, 5, 8, 9]]), [61, 34, 9, 8, 5, 2])
```
